## Assignment-6
##### Total number of points: 40
#### Due date: 11/6/2023 11:59 PM

Before you submit this homework, make sure everything runs as expected. First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All). You can discuss with others regarding the homework but all work must be your own.

Steps to evaluate your solutions:

Step-1: Try on Colab or Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the homework's .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as ".ipynb"

Step-6: Go to lms.rpi.edu and upload your homework at the appropriate link to submit this homework.


#### Please note that for any question in this assignment you will receive points ONLY if your solution passes all the test cases including hidden testcases as well. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management.  

### Please provide your response in clear words as a string. You need to answer this question in order to receive points. 

- 1. Did you work in a group setting to study the concepts you are learning in this class -- Assign variable `groupStudy` with either `yes` or `no`.
- 2. How much time did you spend on this assignment? -- Assign variable `AssignmentTime`.
- 3. How much time did you spend towards learning the concepts taught in this class the past week in terms of hours -- Assign variable `WeeklyHours`. 



<!-- grade_id: Q0, points: 0, grade: false, solution: true -->
```python
#Assign groupstudy; AssignmentTime; WeeklyHours

groupStudy = 'No' #Please assign this variable with yes or no
AssignmentTime = '10' #How much time did you spend w.r.t hours on this assignment
WeeklyHours = '6' #How much time did you spend w.r.t hours the last week learning the concepts taught in our ML class

# YOUR CODE HERE



```

<!-- grade_id: 1, points: 10, grade: true, solution: false -->
```python
#Test cell-1 (10 points)
#DO NOT MODIFY/DELETE THIS CELL 

```

#### Q1 [10 points]. Please make sure you complete the below tasks correctly as the following questions are dependent on this. 

1. Load the SMS data available at the given url which is in the form of a text file. 
    - Please note that this text file contains multiple lines where, each line starts with either "spam" or "ham" followed by a tab ('\t') and the text message ending with a '\n'.
    - The `data` variable contains all these lines.
    - More details about this link could be found here: https://archive.ics.uci.edu/dataset/228/sms+spam+collection 
2. Use the data variable to create `X` and `y` using this dataset where `X` corresponds to a list of all the texts and `y` corresponds to a list of all the class label values. 
    - Make sure you get remove of the new line character for each SMS text in `X`

<!-- grade_id: Q1-sol, points: 0, grade: false, solution: true -->
```python
import numpy as np
import re
url='https://raw.githubusercontent.com/lmanikon/lmanikon.github.io/master/teaching/datasets/SMSSpam.txt'

import requests

response = requests.get(url)
data = response.text
display(data)

X = []
y = []

lines=data.split('\n')
for line in lines:
    if len(line) > 0:
        label, message = line.split('\t')
        X.append(re.sub(r'\n', '', message))
        y.append(label)
print(X)
print(y)

# YOUR CODE HERE


```

<!-- grade_id: 2, points: 4, grade: true, solution: false -->
```python
#Test cell-2 (4 points)
#DO NOT MODIFY/DELETE THIS CELL 

assert len(X)==len(y)
assert len(y)==5574

```

<!-- grade_id: 3, points: 4, grade: true, solution: false -->
```python
#Test cell-3 (4 points)
#DO NOT MODIFY/DELETE THIS CELL 

assert y.count('spam')==747
assert y.count('ham')==4827
```

<!-- grade_id: 4, points: 2, grade: true, solution: false -->
```python
#Test cell-4 (2 points)
#DO NOT MODIFY/DELETE THIS CELL 

```

#### Q2 [10 points]. Preprocessing the SMS text messages -- we will use both `X` and `y` from the previous question

1. Convert the class label column `y` by changing the values of `ham` and `spam` to integer values of 0 and 1 respectively. 
2. Lets now extract numerical features from `X`. Before you work on the following questions, please convert all the text in `X` into lowercase format. 
3. Create a first feature (named `LLen`) that computes the length of each SMS in `X`. 
4. Create a second feature (named `URL`) that contains an integer value of 1 if the corresponding SMS contains a web URL (http or https). IF there is no web url, the appropriate cell value will be integer value of 0.
5. Clean the text by removing all the urls and replacing non-alphanumeric characters with one space. Append these two features (`LLen` and `URL`) along with the TFIDF vectors (and their corresponding vocabulary) that you will extract using TFIDF vectorization approach. 
6. Create a dataframe `X_new` using all the features extracted using `X` and the corresponding feature names should be added. 

<!-- grade_id: Q2-sol, points: 0, grade: false, solution: true -->
```python
from IPython.utils import text
import pandas as pd
import re
from sklearn.feature_extraction.text import TfidfVectorizer

y = [1 if label == 'spam' else 0 for label in y]

X = list(map(str.lower, X))

LLen = list(map(len, X))

URL = []
X_cleaned = []

for text in X:
    if re.search(r'http://|https://', text):
        URL.append(1)
        text = re.sub(r'http\S+', '', text)
    else:
        URL.append(0)

    cleaned_text = re.sub(r'\W+', ' ', text)
    X_cleaned.append(cleaned_text.strip())


vectorizer = TfidfVectorizer()
tfidf_matrix = vectorizer.fit_transform(X_cleaned)
tfidf_features = vectorizer.get_feature_names_out()

tfidf_array = tfidf_matrix.toarray()

LLen_df = pd.DataFrame({'LLen': LLen})
URL_df = pd.DataFrame({'URL': URL})

tfidf_df = pd.DataFrame(tfidf_array, columns=tfidf_features)

X_combined = pd.concat([LLen_df, URL_df, tfidf_df], axis=1)
feature_names = ['LLen', 'URL'] + list(tfidf_features)

X_new = pd.DataFrame(X_combined, columns= feature_names)
print(X_new.columns)

# YOUR CODE HERE

```

<!-- grade_id: 5, points: 4, grade: true, solution: false -->
```python
#Test cell-5 (4 points)
#DO NOT MODIFY/DELETE THIS CELL 

assert y.count(1)==747
assert y.count(0)==4827
```

<!-- grade_id: 6, points: 4, grade: true, solution: false -->
```python
#Test cell-6 (4 points)
#DO NOT MODIFY/DELETE THIS CELL 

assert X_new['URL'].sum()==22
assert round(X_new['LLen'].mean(),2)==80.48
```

<!-- grade_id: 7, points: 2, grade: true, solution: false -->
```python
#Test cell-7 (2 points)
#DO NOT MODIFY/DELETE THIS CELL 

assert len(X_new.columns)==8691
```

#### Q3 [5 points]. Unsupervised learning

1. Perform standardization of `X_new` using StandardScaler function to obtain `X_std`. 
2. Consider the standardized dataframe `X_std` to perform k-means clustering with the value of k as 10 (meaning number of clusters) and `random_state` of 10. 
3. Assign the silhouette score of the clustering to variable named `sil_score`.

<!-- grade_id: Q3-sol, points: 0, grade: false, solution: true -->
```python
#K-means clustering
from sklearn.preprocessing import StandardScaler
from sklearn.cluster import KMeans
from sklearn.metrics import silhouette_score

scaler = StandardScaler()
X_std = scaler.fit_transform(X_new)

kmeans = KMeans(n_clusters=10, random_state=10, n_init=10)
cluster_labels = kmeans.fit_predict(X_std)

sil_score = silhouette_score(X_std, cluster_labels)

print("Silhouette Score:", sil_score)

# YOUR CODE HERE





```

<!-- grade_id: 8, points: 3, grade: true, solution: false -->
```python
#Test cell-8 (7 points)
#DO NOT MODIFY/DELETE THIS CELL 

assert round(X_std.mean())==0
assert round(X_std.std())==1
```

<!-- grade_id: 9, points: 2, grade: true, solution: false -->
```python
#Test cell-9 (3 points)
#DO NOT MODIFY/DELETE THIS CELL 

assert round(sil_score,3)==0.052
```

#### Q4 [5 points]. Supervised learning

1. Consider the dataframe `X_std` and `y` to first perform train_test_split operation with `test_size` as 0.3 and `random_state` as 12. Use the variables `X_train`, `y_train` for training and `X_test`, `y_test` for testing datasets. 
2. Train a K-Nearest Neighbors classifier (initialized as variable `knn`) using these parameters: `n_neighbors`=5 and the `kd_tree` algorithm.
3. Compute the accuracy_score and assign it to variable `acc_score`. 

<!-- grade_id: cell-b13d1e5cfbcc8d22, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from sklearn.model_selection import train_test_split
from sklearn.neighbors import KNeighborsClassifier
from sklearn.metrics import accuracy_score

X_train, X_test, y_train, y_test = train_test_split(X_std, y, test_size=0.3, random_state=12)

knn = KNeighborsClassifier(n_neighbors=5, algorithm='kd_tree')

knn.fit(X_train, y_train)

y_pred = knn.predict(X_test)

acc_score = accuracy_score(y_test, y_pred)

print(acc_score)

```

<!-- grade_id: 10, points: 5, grade: true, solution: false -->
```python
#Test cell-10 (10 points)
#DO NOT MODIFY/DELETE THIS CELL 

assert round(acc_score,2)== 0.88
assert knn.n_neighbors == 5
```