## Assignment-5
##### Total number of points: 50
#### Due date: 10/23/2023 11:59 PM

Before you submit this homework, make sure everything runs as expected. First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All). You can discuss with others regarding the homework but all work must be your own.

Steps to evaluate your solutions:

Step-1: Try on Colab or Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the homework's .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as ".ipynb"

Step-6: Go to lms.rpi.edu and upload your homework at the appropriate link to submit this homework.


#### Please note that for any question in this assignment you will receive points ONLY if your solution passes all the test cases including hidden testcases as well. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management.  

### Please provide your response in clear words as a string. You need to answer this question in order to receive points. 

- 1. Did you work in a group setting to study the concepts you are learning in this class -- Assign variable `groupStudy` with either `yes` or `no`.
- 2. How much time did you spend on this assignment? -- Assign variable `AssignmentTime`.
- 3. How much time did you spend towards learning the concepts taught in this class the past week in terms of hours -- Assign variable `WeeklyHours`. 



<!-- grade_id: Q0, points: 0, grade: false, solution: true -->
```python
#Assign groupstudy; AssignmentTime; WeeklyHours

groupStudy = 'no' #Please assign this variable with yes or no
AssignmentTime = '4' #How much time did you spend w.r.t hours on this assignment
WeeklyHours = '5' #How much time did you spend w.r.t hours the last week learning the concepts taught in our ML class

# YOUR CODE HERE



```

<!-- grade_id: 1, points: 10, grade: true, solution: false -->
```python
#Test cell-1 (10 points)
#DO NOT MODIFY/DELETE THIS CELL 

```

#### Q1 [10 points]. Please make sure you complete the below tasks correctly as the following questions are dependent on this. 

1. Load the California housing data from sklearn library (library is already provided here below). 
    - This dataset was obtained from the StatLib repository https://www.dcc.fc.up.pt/~ltorgo/Regression/cal_housing.html
2. Create a dataframe `X` representing the features -- both data and the feature names should be included. 
3. Save the class labels in an array `y_true`. 
4. The given target labels (`cali.target`) are continuous -- convert them into discrete values -- 1 and 2 using this approach and save them in an array `y_true_cat`.
    - all the values (v1) in cali.target should become 1 if v1>=0.0 and v1<2.0; 
    - all the values (v2) in cali.target should become 2 if v2>=2.0 and v2<5.5;
    - Note that these new/transformed labels 1 and 2 are **integers**
    - Please save these new discrete values in the array `y_true_cat`
    - Original continuous values are in the array `y_true`


<!-- grade_id: Q1, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import fetch_california_housing

#Load the dataset and save it as a dataframe
cali = fetch_california_housing()

feature_names = cali.feature_names
X = pd.DataFrame(cali.data, columns=feature_names)

y_true = cali.target
y_true_cat = []

for v in y_true:
    if v >= 0.0 and v < 2.0:
        y_true_cat.append(1)
    elif v >= 2.0 and v < 5.5:
        y_true_cat.append(2)

y_true_cat = np.array(y_true_cat)
print("Original Continuous Labels (y_true):")
print(y_true[:10])

print("\nDiscrete Labels (y_true_cat):")
print(y_true_cat[:10])
# YOUR CODE HERE


```

Output:
```
Original Continuous Labels (y_true):
[4.526 3.585 3.521 3.413 3.422 2.697 2.992 2.414 2.267 2.611]

Discrete Labels (y_true_cat):
[2 2 2 2 2 2 2 2 2 2]
```

<!-- grade_id: 2, points: 10, grade: true, solution: false -->
```python
#Test cell-2 (10 points)
#DO NOT MODIFY/DELETE THIS CELL 

assert set(y_true_cat)=={1,2}
assert (list(y_true_cat)).count(1)==11885
assert (list(y_true_cat)).count(0)==0
assert len(X)==len(y_true)
assert len(X.columns)==8
```

#### Q2 [10 points]. Linear regression: Use the California housing dataset loaded into the dataframe `X`

1. **Without using train_test_split() function**
    - Use the 1st 11885 rows (i.e., index 0,1,....11884) where the corresponding features are loaded as `train_X` and corresponding labels as `train_y`
    - Use the remaining rows as testing data -- `test_X` and `test_y` for the features and labels respectively.
    - Fit a linear regression line using the training data; then use it to predict labels for testing data as shown in the lecture notebook-14. Please use the default parameters when calling the linear regression function. 
    - Please make sure you call the appropriate libraries.
    - Hint -- Use the original class labels as you expect for Linear Regression. 

2. Measure the mean-squared error (MSE) `mse_split2` using the predicted labels with testing labels.
    - Round the `mse_split2` to 2 values after the decimal point. 
    - Hint -- Check out the libraries given to guess which functions to use to compute MSE.


<!-- grade_id: Q2-sol, points: 0, grade: false, solution: true -->
```python
#Loaded the dataset as X in the previous question 
from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_squared_error

num_train_rows = 11885
train_X = pd.DataFrame(X[:num_train_rows])
train_y = pd.DataFrame(y_true[:num_train_rows])
test_X = pd.DataFrame(X[num_train_rows:])
test_y = pd.DataFrame(y_true[num_train_rows:])

mse_split2=0


regressor = LinearRegression()

regressor.fit(train_X, train_y)

predicted_y = regressor.predict(test_X)

mse_split2 = mean_squared_error(test_y, predicted_y)

mse_split2 = round(mse_split2, 2)

print("Mean Squared Error: ", mse_split2)
# YOUR CODE HERE

```

Output:
```
Mean Squared Error:  0.55
```

<!-- grade_id: 3, points: 6, grade: true, solution: false -->
```python
#Test cell-3 (6 points)
#DO NOT MODIFY/DELETE THIS CELL 

assert len(train_X)==11885
assert len(test_y) == 8755
assert mse_split2 == 0.55

```

<!-- grade_id: 4, points: 4, grade: true, solution: false -->
```python
#Test cell-4 (4 points)
#DO NOT MODIFY/DELETE THIS CELL 

```

#### Q3 [10 points]. Linear regression: Use the California housing dataset loaded into the dataframe `X`

1. **Without using train_test_split() function**
    - Use the 1st 8755 rows (i.e., index 0,1,....8754) where the corresponding features are loaded as `test_X` and corresponding labels as `test_y`
    - Use the remaining rows as testing data -- `train_X` and `train_y` for the features and labels respectively.
    - Fit a linear regression line using the training data; then use it to predict labels for testing data as shown in the lecture notebook-14. Please use the default parameters when calling the linear regression function. 
    - Please make sure you call the appropriate libraries.
    - Hint -- Use the original class labels as you expect for Linear Regression. 

2. Measure the mean-squared error (MSE) `mse_split3` using the predicted labels with testing labels.
    - Round the `mse_split3` to 2 values after the decimal point. 
    - Hint -- Check out the libraries given to guess which functions to use to compute MSE.

3. What does the difference between `mse_split2` and `mse_split3` highlight w.r.t the process of splitting the data into training and testing datasets, if any? Please save your response as a string and assign it to variable `your_response3`. 

<!-- grade_id: Q3-sol, points: 0, grade: false, solution: true -->
```python
#Loaded the dataset as X in 1st question 
from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_squared_error

num_test_rows = 8755
train_X = pd.DataFrame(X[num_test_rows:])
train_y = pd.DataFrame(y_true[num_test_rows:])
test_X = pd.DataFrame(X[:num_test_rows])
test_y = pd.DataFrame(y_true[:num_test_rows])

mse_split3=0

regressor = LinearRegression()

regressor.fit(train_X, train_y)

predicted_y = regressor.predict(test_X)

mse_split3 = mean_squared_error(test_y, predicted_y)

mse_split3 = round(mse_split3, 2)

print("Mean Squared Error:", mse_split3)
your_response3='For mse_split2 we use 11885 rows and for mse_split3 we used 8755 rows for training and testing set which leads to different results'
print(your_response3)
# YOUR CODE HERE

```

Output:
```
Mean Squared Error: 0.56
For mse_split2 we use 11885 rows and for mse_split3 we used 8755 rows for training and testing set which leads to different results
```

<!-- grade_id: 5, points: 6, grade: true, solution: false -->
```python
#Test cell-5 (6 points)
#DO NOT MODIFY/DELETE THIS CELL 

assert len(train_X)==11885
assert len(test_y) == 8755
assert mse_split3 == 0.56

```

<!-- grade_id: 6, points: 4, grade: true, solution: false -->
```python
#Test cell-6 (4 points)
#DO NOT MODIFY/DELETE THIS CELL 

```

#### Q4 [10 points]. Decision Trees: Use the California housing dataset loaded into `X` and use discretized labels as processed in Q1

1. Normalize the data using `MinMaxScaler` function and save it as a variable `X_minmax`
2. Shift the original dimensions into Principal Component (PC) space, where you choose only the top-k components that capture at the max 98% of the variance. Transform the data into this new top-k PC space and save it as `X_pca`.  
3. Initialize a Decision Tree classifier variable using `DecisionTreeClassifier` function with a `random_state=10`.
4. Now using the transformed data, and the classifier variable, compute the average cross validation score when the decision tree classifier is evaluated over a 5-fold cross validation (use `cross_val_score` function). Assign this value to variable `avg_cross_val`. You can use numpy's mean function to compute the average score.

<!-- grade_id: Q4-sol, points: 0, grade: false, solution: true -->
```python
#Using feature set and discretized class labels building a decision tree classifier 
from sklearn.tree import DecisionTreeClassifier
from sklearn.model_selection import cross_val_score
from sklearn.decomposition import PCA
from sklearn.preprocessing import MinMaxScaler
import numpy as np


X_minmax = MinMaxScaler().fit_transform(X)

pca = PCA(n_components=0.97)
X_pca = pca.fit_transform(X_minmax)

clf = DecisionTreeClassifier(random_state=10)

avg_cross_val = np.mean(cross_val_score(clf, X_pca, y_true_cat, cv = 5))

avg_cross_val = round(avg_cross_val, 2)

print("Average Cross Validation Score:", avg_cross_val)



# YOUR CODE HERE

```

Output:
```
Average Cross Validation Score: 0.68
```

<!-- grade_id: 7, points: 8, grade: true, solution: false -->
```python
#Test cell-7 (8 points)
#DO NOT MODIFY/DELETE THIS CELL 
assert round(np.mean(X_minmax[:,0]),2)==0.23
assert round(avg_cross_val, 2)==0.68
assert round(np.mean(X_pca[0]), 2)==0.18
```

<!-- grade_id: 8, points: 2, grade: true, solution: false -->
```python
#Test cell-8 (2 points)
#DO NOT MODIFY/DELETE THIS CELL 

```