## Assignment-4
##### Total number of points: 50
#### Due date: 10/16/2023 11:59 PM

Before you submit this homework, make sure everything runs as expected. First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All). You can discuss with others regarding the homework but all work must be your own.

Steps to evaluate your solutions:

Step-1: Try on Colab or Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the homework's .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as ".ipynb"

Step-6: Go to lms.rpi.edu and upload your homework at the appropriate link to submit this homework.


#### Please note that for any question in this assignment you will receive points ONLY if your solution passes all the test cases including hidden testcases as well. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management.  

### Please provide your response in clear words as a string. You need to answer this question in order to receive points. 

- 1. Did you work in a group setting to study the concepts you are learning in this class -- Assign variable `groupStudy` with either `yes` or `no`.
- 2. How much time did you spend on this assignment? -- Assign variable `AssignmentTime`.
- 3. How much time did you spend towards learning the concepts taught in this class the past week in terms of hours -- Assign variable `WeeklyHours`. 



<!-- grade_id: Q0, points: 0, grade: false, solution: true -->
```python
#Assign groupstudy; AssignmentTime; WeeklyHours

groupStudy = '' #Please assign this variable with yes or no
AssignmentTime = '' #How much time did you spend w.r.t hours on this assignment
WeeklyHours = '' #How much time did you spend w.r.t hours the last week learning the concepts taught in our ML class

# YOUR CODE HERE

```

<!-- grade_id: 1, points: 10, grade: true, solution: false -->
```python
#Test cell-1 (10 points)
#DO NOT MODIFY/DELETE THIS CELL 

```

We will be using the popular breast cancer wisconsin dataset for this homework (more details here: [url](https://scikit-learn.org/stable/modules/generated/sklearn.datasets.load_breast_cancer.html))

#### Q1. Building and preparing a dataset for model generation. 

1. Load the data separately as `X` and `y_true` corresponding to features and class label or target variable respectively. WHen you build the feature set `X` make sure `X` is a pandas dataframe that contains both data and feature names. `y_true` could be a pandas Series. <br><br>

2. Create a dataframe `df` using both `X` and `y_true` where, you will be using the feature names exactly as they are in `X`. <br><br>

3. Create a new feature column named `y_true_cat` in `df` where, you convert the numerical values of 0 and 1 to categorical values.  Use the appropriate mapping to ensure in such a way that there are a total of 212 `Malignant` and 357 `Benign` instances in `y_true_cat`. <br><br>

4. Create a new feature column (integer type) called `y_pred` where, the first 300 values are `0`s, and the remaining 269 values are `1`s. <br><br>

5. Now use both columns `y_true` and `y_pred` as the ground truth labels and predicted labels to compute Precision (`q1_prec`), Recall (`q1_rec`) and Accuracy (`q1_acc`).

<!-- grade_id: Q1-sol, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import load_breast_cancer

#Load the dataset and save it as a dataframe
bcancer = load_breast_cancer()

# YOUR CODE HERE




```

<!-- grade_id: 2, points: 5, grade: true, solution: false -->
```python
#Test cell-2 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
assert df['y_true_cat'].value_counts()['Benign']==357
assert df['y_true_cat'].value_counts()['Malignant']==212
```

<!-- grade_id: 3, points: 5, grade: true, solution: false -->
```python
#Test cell-3 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
assert 'y_true' in df.columns
assert len(X)==569
assert len(df.columns)==33
```

<!-- grade_id: 4, points: 5, grade: true, solution: false -->
```python
#Test cell-4 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
```

#### Q2. Data preprocessing and model evaluation. 


1. Drop the columns `y_true_cat` and `y_pred` from the dataframe `df`. <br><br>

2. Create `X` and `y` from the dataframe `df` where, `X` corresponds to only the feature columns, and `y` corresponds to the class label column `y_true`. Please make sure you check the casing -- `X` and `y` are in uppercase and lowercase formats respectively.  <br><br>

3. Split the data into training (75%) and testing (25%) data using only the parameters X (feature set) and y (class label `y_true`, `test_size`, and `random_state`=10. Dataframe `df` and `y_true` were initialized in the previous question. This generates these four entities as outputs (please use the same variable names as follows): 
    - `X_train`: Training feature columns
    - `X_test`: Testing feature columns
    - `y_train`: Training labels
    - `y_test`: Testing labels
<br><br>
4. Please copy and paste this code in your code cell after step 3 exactly as it is and execute it to compute the predicted class labels `y_pred`. 

`from sklearn.tree import DecisionTreeClassifier 
 clf = DecisionTreeClassifier(random_state=0) 
 clf.fit(X_train, y_train) 
 y_pred = clf.predict(X_test)`

<br>
5. Finally, compute the precision, recall, and accuracy scores assigned to the variables `q2_prec`, `q2_rec`, and `q2_acc` respectively. 


<!-- grade_id: Q2-sol, points: 0, grade: false, solution: true -->
```python
#Split the data into training (75%) and testing (25%) data

#X_train: Training feature columns
#X_test: Testing feature columns
#y_train: Training labels
#y_test: Testing labels

# YOUR CODE HERE



```

<!-- grade_id: 5, points: 5, grade: true, solution: false -->
```python
#Test cell-5 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
assert 'y_true' in df.columns
assert df.isnull().values.sum()==0
assert set(y)=={0, 1}
assert len(X_train)==426
```

<!-- grade_id: 6, points: 5, grade: true, solution: false -->
```python
#Test cell-6 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
```

#### Q3. Data preprocessing and dimensionality reduction using PCA. 


1. Consider the same `df` from the previous question (Q2 -- after modifications). Create `X` and `y` from the dataframe `df` where, `X` corresponds to only the feature columns, and `y` corresponds to the class label column `y_true`. Please make sure you check the casing -- `X` and `y` are in uppercase and lowercase formats respectively. 

2. Perform standardization using StandardScaler function and save its output as `X_std`. 

3. Consider only the top-`k` Principal components that that will retain not more than 78% of the information present in the original dataset. In other words just choose the first `k` principal components that capture 78% of variance only. Please assign the integer variable `k` with the total number of these top principal components. 

4. Translate the data only into those `k` new feature (which are the principal components) space and save the transformed data as `X_pca`. 

<!-- grade_id: Q3-sol, points: 0, grade: false, solution: true -->
```python
#Perform standardization using StandardScaler function and perform PCA.

#Consider only the top-k Principal components that that will retain not more than 78% of the information present in the original dataset.

#In other words just choose the first k principal components that capture 78% of variance only.

#How many such PCs should be considered?

import numpy as np

# YOUR CODE HERE




```

<!-- grade_id: 7, points: 5, grade: true, solution: false -->
```python
#Test cell-7 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
assert len(X)==569
assert len(X.columns)==30
assert len(X)==len(y)
```

<!-- grade_id: 8, points: 5, grade: true, solution: false -->
```python
#Test cell-8 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
assert round(np.mean(X_std),2)==0
assert round(np.std(X_std),2)==1
assert len(X_pca[0])==k
assert k<len(X.columns)
```

<!-- grade_id: 9, points: 5, grade: true, solution: false -->
```python
#Test cell-9 (5 points)
#DO NOT MODIFY/DELETE THIS CELL
```